<a href="https://colab.research.google.com/github/hieutrinhds/learn_git/blob/master/trinh_trung_hieu.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Hi Philippines, firstly, please enter your exam username to the variable STUDENT_NAME by selecting one of the exam username in the below list!

```python
# for example, your username is huynh_nguyen_minh_thong
STUDENT_NAME = "trinh_trung_hieu"
```

Secondly, please rename your exam jupyter notebook filename to be your username like this: huynh_nguyen_minh_thong.ipynb

Please just choose your correct exam username above using one of those names below so my script will automatically detect one of the names below then generate score and feedback for the correct person:

thinh_tran

beka_martinez

trinh_tuan_anh

huy_nguyen

minh_vu

abi_vu

trong_hieu_nguyen

trinh_trung_hieu

hung_nham

ai_nguyen

hoa_han

hong_ke_kien

hoang_thai_cong

nguyen_thong

dung_nguyen

hieu_pham

---

# Welcome to the first test of Philippines Machine Learning class!

There are a **20 questions** of multiple questions and coding challenges. Each of them will worth some points. Please **follow the strict output format** carefully because we will use a hidden testing script to automatically mark your tests when you submit. There will be hidden test cases that only we know so make sure your coding sections cover all potential cases.

*Total score is 100 points (20 main questions) + 20 points(one bonus question).*

Note: this is an **open book test** so feel free to google or refer back to your learning notebooks. 

<img src="https://course_report_production.s3.amazonaws.com/rich/rich_files/rich_files/5741/s300/logo-cs-01-01.png" width="300">



---
## Multiple choices with multiple answers
Remember your answer always will be in **list**. **There will be maybe multiple correct answers or just single answer**.

1) choice 1

2) choice 2

3) choice 3

4) choice 4

5) choice 5


Example - Your answer **must** be like **one** of these formats like below:

return [1]

OR

return [1,3]

OR

return [2,3,4]


---
### Question 0a (0 point) Example Question 1


Which keyword in **Python** to define a function?

1) function

2) def

3) @

4) Class

5) key


```python
def question_0a():
  # your answer should be in this format for single answer (you choose answer 2 in this case)
  return [2]
```

```python
# This is how the hidden test script works in the hidden cell (you don't need to worry about this part during the exam)
# Make sure your answer pass the test script with no error output (this is how our public and hidden test script works)
assert sorted(question_0a()) == [2]
```

---
### Question 0b (0 point) Example Question 2


Which ones are keywords in Python?

1) test

2) if

3) return

4) break

5) yes


```python
def question_0b():
  # your answer should be in this format for multiple answers (you choose answer 2, 3 and 4 in this case)
  return [2,3,4]
```

```python
# This is how the hidden test script works in the hidden cell (you don't need to worry about this part during the exam)
# Make sure your answer pass the test script with no error output (this is how our public and hidden test script works)
assert sorted(question_0b()) == [2,3,4]
```

---
### Question 1 (2 point)

```python
name = "Minh"
height = 180
```
Which is the valid ways to format string mixed with variable in **Python** but not Jinja ?
```python
1) "My name is " + name + " and my height is " + height
2) "My name is {} and my height is {}".format(name, height)
3) f"My name is {name} and my height is {height}"
4) "My name is %s and my height is %s" %(name, height)
5) "My name is {{name}} and my height is {{height}}"
```


```python
name = "Minh"
height = 180
print("My name is %s and my height is %s" %(name, height))
```

Output:
```
My name is Minh and my height is 180
```

<!-- grade_id: question_1, points: 0, grade: false, solution: true -->
```python
def question_1():
  # YOUR ANSWER HERE
  return [2, 3]
```

<!-- grade_id: question_1_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 2 (2 point)

If the function doensn't have "return" to output or give back any value, what is the default output when calling that function  ?

1) Int

2) Bool

3) Void

4) None



<!-- grade_id: question_2, points: 0, grade: false, solution: true -->
```python
def question_2():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_2_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 3 (2 point)

In Python, what is the correct math operation orders of precedence (from most prioritised to least prioritised) ?

*Given that:*

a) Parentheses

b) Exponential

c) Multiplication

d) Division

e) Addition

f) Subtraction

*Which order?*

1) a,b,c,d,e,f

2) b,a,c,d,e,f

3) b,a,d,c,e,f

4) a,b,c,d,f,e



<!-- grade_id: question_3, points: 0, grade: false, solution: true -->
```python
def question_3():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_3_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 4 (2 point)

```python
a = "Hello World Philippines!"
b = a[::]
c = b[::-1]
d = a[:]
e = a[:3]
g = a[0:-1:2]
h = a[1:-2:-1]
j = a[-1:0:-1]
l = a[-3:-1]
p = a[2:-2:-1]
```

*Which one give you empty string or invalid or error?*

1) g,j

2) h,l

3) p,h

4) g,b



<!-- grade_id: question_4, points: 0, grade: false, solution: true -->
```python
def question_4():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_4_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 5 (2 point)

In the previous question, why that slicing gives empty string or invalid or error?

*Which reason?*

1) invalid syntax

2) the start position and end position is crossed due to the negative step interval

3) do not have any information/argument for slicing

4) string memory is corrupted due to many time accesses and slicings.



<!-- grade_id: question_5, points: 0, grade: false, solution: true -->
```python
def question_5():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_5_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 6 (2 point)
```python
w="good morning Philippines"
v=('a','o','e','u','i')
l="aoeui"
u=list(l)
```
Which list comprehension return list of vowel characters in the string like below output?
['o', 'o', 'o', 'i', 'o', 'a']

Which list comprehension?

```python
1) [w for w in w if w in l]
2) [t for t in w if t in u]
3) [r for r in w if r in v]
4) [l for v in w if x in w]
```




<!-- grade_id: question_6, points: 0, grade: false, solution: true -->
```python
def question_6():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_6_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 7 (2 point)
Remember we have learnt about dictionary in python, we know every entry in a dictionary will have a key and a value.

Which one is not true about dictionary key?

1) There is no key which is the same

2) Key must be immutable

3) Key can be string

4) Key must be integer

5) Value of dictionary can be any data type

<!-- grade_id: question_7, points: 0, grade: false, solution: true -->
```python
def question_7():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_7_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 8 (2 point)
Which git commands creates a copy of repository on the localhost/workspace?

1) commit

2) pull

3) merge

4) clone

5) push

<!-- grade_id: question_8, points: 0, grade: false, solution: true -->
```python
def question_8():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_8_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 9 (2 point)
In your weekly project, you already have a local repository folder, and other teammate(s) pushed their updates to the remote repository on github. which git command you need to use to download those changes into your local folder?

1) checkout

2) pull

3) merge

4) clone

5) push

<!-- grade_id: question_9, points: 0, grade: false, solution: true -->
```python
def question_9():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_9_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 10 (2 point)
For your weekly project, you created a project folder, initialised a new git repository by doing "git init" and created a new HTML file named "index.html", which following git command will not work or return error warnings if running it?

1) git status

2) git add index.html

3) git add .

4) git commit -m "homepage html file added"

5) git push

<!-- grade_id: question_10, points: 0, grade: false, solution: true -->
```python
def question_10():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_10_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 11 (2 point)
For bash commands, we have definitions as below:

a) **pwd** is to print out current working directory

b) **ls** is to list down all file(s)/folder(s) in current folder

c) **cd** is to change current directory to another folder

d) **mkdir** is to make a new directory

e) **rm** is to remove folder(s)/file(s)

f) **cp** is to copy file(s)/folder(s)

g) **mv** is to move file(s)/folder(s)

h) **touch** is to create empty file(s)

l) **cat** is to show the content of the file

o) **grep** for searching

p) **sed** for transforming

What one is incorrected?

1) c, g, p

2) e, o, h

3) o, l, e

4) None of them

<!-- grade_id: question_11, points: 0, grade: false, solution: true -->
```python
def question_11():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_11_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 12 (2 point)

What is incorrect about PRIMARY KEY constraint in database?

1) It uniquely identifies each row/record in table

2) It must have unique value

3) It can either on one column or multiple columns

4) We need to always manually define the value of primary key when inserting a new row/record

5) It cannot contain NULL value


<!-- grade_id: question_12, points: 0, grade: false, solution: true -->
```python
def question_12():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_12_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 13 (2 point)

Which are possible relationship types between two tables in relational database?

1) One-to-One relationship

2) Few-to-One relationship

3) One-to-Many relationship

4) Many-to-Single relationship

5) Many-to-Many relationship


<!-- grade_id: question_13, points: 0, grade: false, solution: true -->
```python
def question_13():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_13_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
### Question 14 (2 point)

What types of join is available in relational database?

1) Inner Join

2) Outside Join

3) Down Join

4) Left Join

5) Right Join

6) Up Join

7) Full (Outer) Join


<!-- grade_id: question_14, points: 0, grade: false, solution: true -->
```python
def question_14():
  # YOUR ANSWER HERE
  return 
```

<!-- grade_id: question_14_answer, points: 2, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
```

---
## Code Challenges

Each question has description of what the function suppose to return. The function name and declaration is done for you. **Please do not change the function name and its input argument(s)**. Therefore, only complete the function body and return **correct datatype of value**.

**Remember that your part is to only edit/modify where there is "YOUR ANSWER HERE" comment section.**

Since we use functions, so **please do not use print to return values, only use "return" keyword to return final value!**

**Have fun and enjoy coding section!**


---
### Question 15 (5 points)
Complete the function to remove duplicated elements in the list and return the sorted result

``` python
remove_duplicate([2,3,2,5,6,7,10,6,3])

[2,3,5,6,7,10]
```


<!-- grade_id: question_15, points: 0, grade: false, solution: true -->
```python
def remove_duplicate(dup):
  # YOUR ANSWER HERE
  return 
  

print(remove_duplicate([2,3,2,5,6,7,10,6,3]))
```

<!-- grade_id: question_15_answer, points: 5, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
# Make sure your answer pass the test script with no error output (we will have more hidden test cases)
assert remove_duplicate([2,3,2,5,6,7,10,6,3]) == [2, 3, 5, 6, 7, 10]
```

---
### Question 16  (8 points)
Complete the function to calculate the total cost of the receipt input as a dictionary as below

``` python

receipt = {
  "milk": {"weight": 1, "cost": 30},
  "rice": {"weight": 2, "cost": 20},
  "cookie": {"weight": 3, "cost": 30},
  "sugar": {"weight": 1, "cost": 40},
}

process_receipt(receipt)

"Total cost is $120 and weight is 7 kg"
```


<!-- grade_id: question_16, points: 0, grade: false, solution: true -->
```python
receipt = {
  "milk": {"weight": 1, "cost": 30},
  "rice": {"weight": 2, "cost": 20},
  "cookie": {"weight": 3, "cost": 30},
  "sugar": {"weight": 1, "cost": 40},
}

def process_receipt(receipt):
  # YOUR ANSWER HERE
  return 

process_receipt(receipt)
```

<!-- grade_id: question_16_answer, points: 8, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
# Make sure your answer pass the test script with no error output (we will have more hidden test cases)
receipt = {
  "milk": {"weight": 1, "cost": 30},
  "rice": {"weight": 2, "cost": 20},
  "cookie": {"weight": 3, "cost": 30},
  "sugar": {"weight": 1, "cost": 40},
}
assert process_receipt(receipt) == 'Total cost is $120 and weight is 7 kg'
```

---
### Question 17  (12 point)
Complete the function to calculate the possible number of ways to climb to the stair with n steps.

There is **two way** to climb:

**1) Step up one step (+1 step)**

**2) Jump to skip one step (+2 steps)**

Note: Jumping is always +2 steps and you can do as many jumps as you want. Also, you cannot over jump more than n steps like if n = 5, and you are at t=4 then you cannot choose to jump but step up one step only.

**For example**
If n = 5 which means the stairs got 5 steps to reach the top, you are at step 1 (t=1), you have two decisions: 

1) Step up one step (+1) -> t=2

2) Jump to skip one step (+2) -> t=3

When you are t=2 or t=3, you continue to choose what to do next until you reach t=5.


So in total:

1) t=1 -> t=2 -> t=3 -> t=4 -> t=5 (always step up normally) (Red color)

OR

2) t=1 -> t=2 -> t=3 -> t=5 (jump one time - last step) (Green color)

OR

3) t=1 -> t=2 -> t = 4 -> t=5 (jump one time - second step) (Grey color)

OR

4) t=1 -> t=3 -> t=4 -> t=5 (jump one time - first step) (Blue color)

OR

5) t=1 -> t=3 -> t=5 (always jump) (Dark Red color)

Here is my best to sketch the solution :) No shame please!

<img src="https://i.imgur.com/FdX5qnj.png" width="400">

**In total, there is 5 ways (5 color paths) to climb up stair n = 5**. In this case, the *number of ways* is *coincident* to be the same with **n** for n=5. Usually number of ways is way bigger for big n.

Hints: it is easier to do it in a **recursive** way!


``` python
# with t=1 (starting step) and n=5 (highest steps in the stair)
climb_stairs(1, 5)

5
```


<!-- grade_id: question_17, points: 0, grade: false, solution: true -->
```python
count = 0
def climb_stair(t,n):
  # YOUR ANSWER HERE
  return 

climb_stair(1,5) 
```

<!-- grade_id: question_17_answer, points: 12, grade: true, solution: false -->
```python
# Do not delete this cell, this contains marking script! (If you accidentally do it, then undo delete cell Ctrl+M Z)
# Make sure your answer pass the test script with no error output (we will have more hidden test cases)
count = 0
assert climb_stair(1,4) == 3

count = 0
assert climb_stair(1,5) == 5

count = 0
assert climb_stair(1,7) == 13
```

---
### Question 18 (15 point)
Scrapping a web link: https://tuoitrenews.vn/international

Return a list of dictionary of article names and descriptions of **"Latest News"** section at bottom left of the page (like only 10 articles there)! The format should be exactly like below!

**The function must return a tuple of two elements**:

1) String of "Number of latest articles is ..."

2) List of dictionary of articles where each article has title and description as keys as below.

``` python
scrape_monthly_startup_fr()

('Number of latest articles is 10',
 [{'description': "'It is obvious that Xi could not have visited Wuhan earlier because the risk of him contracting the virus there was too high initially'",
   'title': "China's President Xi visits Wuhan as number of new coronavirus cases tumbles"},
  {'description': 'The daily rate of new infections in South Korea fell to its lowest level in 11 days on Monday',
   'title': 'South Korea reports 35 new coronavirus cases, total at 7,513'},
...
  {'description': 'All museums, gyms, cultural centers, ski resorts and swimming pools will be shut in the targeted zones, while leave was canceled for health workers as Italy’s hospitals sag under the pressure of the virus',
   'title': 'Italy orders lockdown of swathes of north to try to stop coronavirus'}])
```

**Note: there is no hidden test script for this question since news content is changing everyday so this question we will mark the answer manually.**


<!-- grade_id: question_18, points: 15, grade: true, solution: true -->
```python
import requests
from bs4 import BeautifulSoup
def scrape_monthly_startup_fr():
  # YOUR ANSWER HERE
  return 

scrape_monthly_startup_fr()
```

---
### Question 19 (12 point)
We have learnt decorators and this challenge is focused on the flow of stacked decorators.

We have two decorators attached to the function fancy_print below to decorate the output string (such a nice pun). 

Complete the body of k function and l function of print out top and bottom strings as below!

``` python
fancy_print("Philippines 2020")

##########
----------
Philippines 2020
----------
##########
```

**Note: there is no hidden test script for this question so we will mark the answer manually. Let you guys be creative here!**

<!-- grade_id: question_19, points: 12, grade: true, solution: true -->
```python
def g(func_input):
    """
    This decorator only needs to take care of "#" part
    """
    def k(*args, **kwargs):
      # YOUR ANSWER HERE
      return 
      func_input(*args, **kwargs)
      # YOUR ANSWER HERE
      return 
    return k
def h(func_input):
    """
    This decorator only needs to take care of "-" part
    """
    def l(*args, **kwargs):
      # YOUR ANSWER HERE
      return 
      func_input(*args, **kwargs)
      # YOUR ANSWER HERE
      return 
    return l

@g  # this decorator run first
@h  # this decorator run later
def fancy_print(m):
    print(m)

fancy_print("Philippines 2020")
```

---
### Question 20 (15 point)

In the random.db, we have Items table, Sizes table and Manufacturers table. Try to **do three way join to get the result table as below with the condition is that the price is larger than 10 and sort the result by item name in descending order** :)


**Items**

| id | name   | price   | size_id   | manufacturer_id
|------|------|------|------|------|
|   1  | chair| 10| 3|1
|   2  | table| 20| 3|1
|   3  | cup| 30| 2|2
|   4  | glass| 16| 1|2
|   5  | lamp| 20| 1|3
|   6  | door| 25| 2|3


**Sizes**

| id | name   |
|------|------|
|   1  | small|
|   2  | medium|
|   3  | big|


**Manufacturers**

| id | name   |
|------|------|
|   1  | Honda|
|   2  | Ikea|
|   3  | Suyuki|



**Result Table**

|id|name|price|size |manufacturer
|------|------|------|------|------|
|2|	table|	20|	big|	Honda
|5	|lamp|	20|	small|	Suyuki
|4	|glass|	16|	small|	Ikea
|6	|door|	25|	medium|	Suyuki


**Note: there is no hidden test script for this question so we will mark the answer manually for each student.**

```python
# download the random.db file for our database
!wget -q 'https://raw.githubusercontent.com/dhminh1024/practice_datasets/master/random.db'
```

```python
# connect to sqlite db file (random.db)
import sqlite3
conn = sqlite3.connect('random.db')
```

<!-- grade_id: question_20, points: 20, grade: true, solution: true -->
```python
import pandas as pd
# you can use pd.read_sql_query to query SQL and display the table like above
# YOUR ANSWER HERE
return 
```

---
# Bonus Final Fun Challenge - Word Cloud (20 point)


<img src="https://i.imgur.com/TRGcH6z.png" width="300">

For this project, you'll create a "word cloud" from a text by writing a script.  This script needs to process the text, remove punctuation/non-alphabets, ignore case and words that do not contain all alphabets, count the frequencies, and ignore uninteresting or irrelevant words.  A dictionary is the output of the `calculate_frequencies` function.  The `wordcloud` module will then generate the image from your dictionary.

```python
# Here are all the installs and imports you will need for your word cloud script

!pip install wordcloud

import wordcloud
import numpy as np
from matplotlib import pyplot as plt
from IPython.display import display
import requests
```

Whew! That was a lot. All of the installs and imports for your word cloud script have been completed. 
<br><br>
**IMPORTANT!** If this was your first time running the above cell containing the installs and imports, you will need save this notebook now. Then under the File menu above,  select Close and Halt. When the notebook has completely shut down, reopen it. This is the only way the necessary changes will take affect.
<br><br>


```python
# link of text file of Shakespeare books
textlink = 'https://raw.githubusercontent.com/bbejeck/hadoop-algorithms/master/src/shakespeare.txt'
resp = requests.get(textlink)
whole_text=resp.content.decode("utf-8") #decode from byte to string 
whole_text[:1000]
```

As you can see those non-alphabet characeters such as"\ufeff", "\n", "*", "-" in the text. Those characters are not meaningful to us so we need to clean the data so that we replace those character (non-alphabets) with empty string while keeping alphabet characters.


<!-- grade_id: question_21a_answer, points: 10, grade: true, solution: true -->
```python
# remove all non-alphabet characters by keep only alphabet characters, exclude !()-[]{};:'"\,<>./?@#$%^&*_~ punctations
# hopefully the after this code, the whole text will contain alphabet characters and space.
import re
# YOUR ANSWER HERE
return 
whole_text[:1000]
```

Write a function in the cell below that iterates through the words in *whole_text*, removes punctuation, and counts the frequency of each word.  Oh, and be sure to make it ignore word case, words that do not contain all alphabets and boring words like "and" or "the".  Then use it in the `generate_from_frequencies` function to generate your very own word cloud!
<br><br>
**Hint:** Try storing the results of your iteration in a dictionary before passing them into wordcloud via the `generate_from_frequencies` function.

<!-- grade_id: question_21b_answer, points: 10, grade: true, solution: true -->
```python
def calculate_frequencies(file_contents):
    # Here is a list of uninteresting words you can use to process your text
    uninteresting_words = ["the", "a", "to", "if", "is", "it", "of", "and", "or", "an", "as", "i", "me", "my", \
    "we", "our", "ours", "you", "your", "yours", "he", "she", "him", "his", "her", "hers", "its", "they", "them", \
    "their", "what", "which", "who", "whom", "this", "that", "am", "are", "was", "were", "be", "been", "being", \
    "have", "has", "had", "do", "does", "did", "but", "at", "by", "with", "from", "here", "when", "where", "how", \
    "all", "any", "both", "each", "few", "more", "some", "such", "no", "nor", "too", "very", "can", "will", "just", \
    "in", "not", "for", "should", "would", "so", "shall", "on", "thou", "thee", "thy", "than", "s","d", "o", "ll"]
    
    # result is a dictionary which store the frequency of each words in the text 
    # so we need result with word as key and frequency as value
    result = {}
    # YOUR ANSWER HERE
    return 
    
    #wordcloud
    cloud = wordcloud.WordCloud()
    cloud.generate_from_frequencies(result)
    return cloud.to_array()
```

If you have done everything correctly, your word cloud image should appear after running the cell below.  Fingers crossed!

```python
# Display your wordcloud image

myimage = calculate_frequencies(whole_text)
plt.imshow(myimage, interpolation = 'nearest')
plt.axis('off')
plt.show()
```

If your word cloud image did not appear, go back and rework your `calculate_frequencies` function until you get the desired output.  Definitely check that you passed your frequecy count dictionary into the `generate_from_frequencies` function of `wordcloud`. Once you have correctly displayed your word cloud image, you are all done with this project. Nice work!

Note: Make sure you got similar result I got above to receive full marks. In fact, I can see king, lord, good, sir, let, now, man, upon, well and love are popular words in Shakespear text, which make sense.

Note: there is no hidden test script for this question so we will mark the answer manually.

```python

```